---
rule_files:
  - prometheus-alerts.yaml

tests:
  - interval: 5m
    input_series:
      - series: 'celery_task_failed_total{job="celery-exporter", name="test-task", namespace="staging"}'
        values: '1+10x10'
      - series: 'celery_task_suceeded_total{job="celery-exporter", name="test-task", namespace="staging"}'
        values: '1+1x10'
    alert_rule_test:
      - eval_time: 20m
        alertname: CeleryTaskHighFailRate
        exp_alerts:
          - exp_labels:
              job: celery-exporter
              severity: warning
              name: test-task
              namespace: staging
            exp_annotations:
              summary: 'Celery high task fail rate.'
              description: 'More than 5% tasks failed for the task staging/test-task the past 10m.'
              dashboard_url: 'https://grafana.com/d/celery-tasks-by-task-32s3/celery-tasks-by-task?&var-task=test-task'
  - interval: 1m
    input_series:
      - series: 'celery_worker_up{job="celery", hostname="down"}'
        values: '1 0'
      - series: 'celery_worker_up{job="celery", hostname="up"}'
        values: '1 1'
    alert_rule_test:
      - eval_time: 5m
        alertname: CeleryWorkerDown
        exp_alerts:
          - exp_labels:
              job: celery
              severity: warning
              hostname: down
            exp_annotations:
              summary: 'A Celery worker is offline.'
              description: 'The Celery worker down is offline.'
